Gliogenesis from the subventricular zone modulates the extracellular matrix at the glial scar after brain ischemia

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public review):

Summary:

The authors show that SVZ derived astrocytes respond to a middle carotid artery occlusion

(MCAO) hypoxia lesion by secreting and modulating hyaluronan at the edge of the lesion (penumbra) and that hyaluronan is a chemoattractant to SVZ astrocytes. They use lineage tracing of SVZ cells to determine their origin. They also find that SVZ derived astrocytes express Thbs-4 but astrocytes at the MCAO-induced scar do not. Also, they demonstrate that decreased HA in the SVZ is correlated with gliogenesis. While much of the paper is descriptive/correlative they do overexpress Hyaluronan synthase 2 via viral vectors and show this is sufficient to recruit astrocytes to the injury. Interestingly, astrocytes preferred to migrate to the MCAO than to the region of overexpressed HAS2.

Strengths:

The field has largely ignored the gliogenic response of the SVZ, especially with regards to astrocytic function. These cells and especially newborn cells may provide support for regeneration. Emigrated cells from the SVZ have been shown to be neuroprotective via creating pro-survival environments, but their expression and deposition of beneficial extracellular matrix molecules is poorly understood. Therefore, this study is timely and important. The paper is very well written and the flow of results logical.

Comments on revised version:

The authors have addressed my points and the paper is much improved. Here are the salient remaining issues that I suggest be addressed.

We appreciate the feedback by the reviewer, and we are glad that the paper is considered to be much improved. We have done our best to address the remaining issues in this 2nd revision.

The authors have still not shown, using loss of function studies, that Hyaluronan is necessary for SVZ astrogenesis and or migration to MCAO lesions.

This is true. Unfortunately, complete removal of hyaluronan (via Hyase) triggers epilepsy, already described in 1963 by James Young (Exp Neurol paper). Degradation by Hyase also provokes neuroinflammation (Soria et al., 2020 Nat Commun). Two alternatives could be 1) partial depletion with Has inhibitor 4MU (but it is also associated with increased inflammation) or 2) a Has-KO mouse, such as Has3-/- (Arranz et al., 2014), although, to our knowledge, this mouse line is not openly available. We have added a sentence in line 332 addressing this shortcoming: “Loss-of-function studies, using HA-depletion models or HA synthase (Has)deficient mice are still needed to corroborate this finding, although the inflammation associated with HA deficiency might confound interpretation.”

(1) The co-expression of EGFr with Thbs4 and the literature examination is useful.

We thank the reviewer for the kind comment.

(2) Too bad they cannot explain the lack of effect of the MCAO on type C cells. The comparison with kainate-induced epilepsy in the hippocampus may or may not be relevant.

As stated in the previous response, we also found this interesting, and it does warrant further exploration by looking into possible direct NSC-astrocyte differentiation. But we believe that both this possible direct differentiation and the reactive status for these astrocytes are out of the scope of the study. We will not speculate about this in the discussion, either.

(3) Thanks for including the orthogonal confocal views in Fig S6D.

(4) The statement that "BrdU+/Thbs4+ cells mostly in the dorsal area" and therefore they mostly focused on that region is strange. Figure 8 clearly shows Thbs4 staining all along the striatal SVZ. Do they mean the dorsal segment of the striatal SVZ or the subcallosal SVZ? Fig. 4b and Fig 4f clearly show the "subcallosal" area as the one analysed but other figures show the dorsal striatal region (Fig. 2a). This is important because of the well-known embryological and neurogenic differences between the regions.

While it is true that Thbs4 is also expressed in the other subregions of the SVZ (lateral, ventral and medial), as observed in Fig 8. we chose the dorsal area because it is the subregion where we observed the larger increase in slow proliferative NSCs (Thbs4/GFAP/BrdU-positive cells) after MCAO (Fig S3). As observed in the quantifications in Fig S3, we found Thbs4/GFAP/BrdUpositive cells increase in lateral, medial and ventral SVZ, but it is not significant. Therefore, from Fig 4 onwards, we focused on the dorsal SVZ, which the reviewer mentions as “subcallosal” area. We chose the term “dorsal” as stated in single-cell studies (Cebrian-Silla et al, 2021, eLife; Marcy et al., 2023, Sci Adv) and reviews (Sequerra 2014 Front Cell Neurosci) that investigate or mention this subregion, respectively. In the abstract, we are perfectly clear stating that newborn astrocytes migrate frm both dorsal and medial areas.

In Fig 2a, the immunofluorescence image shows medial and lateral SVZ, but at this point in the paper, we have not yet made specific subregional quantifications, and the Nestin, DCX and Thbs4 quantifications refer to the SVZ as a whole, both in the IF and in the WB (Fig 2e-g). We apologize for the confusion. We have clarified this in the text (line 119).

(5) It is good to know that the harsh MCAO's had already been excluded.

(6) Sorry for the lack of clarity - in addition to Thbs4, I was referring to mouse versus rat Hyaluronan degradation genes (Hyal1, Hyal2 and Hyal3) and hyaluronan synthase genes (HAS1 and HAS2) in order to address the overall species differences in hyaluronan biology thus justifying the "shift" from mouse to rat. You examine these in the (weirdly positioned) Fig. 8h,i. Please add a few sentences on mouse vs rat Thbs4 and Hyaluronan relevant genes.

We thank the reviewer for these remarks. We have now added a sentence pointing to the similar internalization and degradation in rat and mouse (reviewed by Sherman et al., 2015). This correction is in line 233. Hyaluronan is, in evolutionary terms, a very “old” molecule, part of the “ancient” glycan-based matrix, before the evolution of proteoglycans and fibrous proteins such as collagen, laminin etc. Hence, its machinery is highly conserved across species.

We have also reorganized the panels in Fig 8, where 8h and 8i were indeed weirdly positioned. We hope that the new version of this figure is more easily readable.

(7) Thank you for the better justification of using the naked mole rat HA synthase.

Reviewer #3 (Public review):

Summary:

The authors aimed to study the activation of gliogenesis and the role of newborn astrocytes in a post-ischemic scenario. Combining immunofluorescence, BrdU-tracing and genetic cellular labelling, they tracked the migration of newborn astrocytes (expressing Thbs4) and found that Thbs4-positive astrocytes modulate the extracellular matrix at the lesion border by synthesis but also degradation of hyaluronan. Their results point to a relevant function of SVZ newborn astrocytes in the modulation of the glial scar after brain ischemia. This work's major strength is the fact that it is tackling the function of SVZ newborn astrocytes, whose role is undisclosed so far.

Strengths:

The article is innovative, of good quality, and clearly written, with properly described Materials and Methods, data analysis and presentation. In general, the methods are designed properly to answer the main question of the authors, being a major strength. Interpretation of the data is also in general well done, with results supporting the main conclusions of this article.

In this revised version, the points raised/weaknesses were clarified and discussed in the article.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Minor points:

(1) Thanks for the clarification.

(2) Thanks for the clarification.

(3) The magnification is not apparent in Fig. 5.

We had removed two brain slices (from 4 to 2) in order to increase the size of the image 2-fold. We have now further increased the TTC panel, 25% from the revised version, 125% from the original.

(4) Thanks for the clarification.

(5) Thanks for the clarification.

(6) Thanks for the clarification.

(7) Thanks for the clarification.

(8) Thanks for the clarification.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Activation of the subventricular zone (SVZ) following cerebral ischemia is one of the brain’s early responses to counteract neuron loss and minimize tissue damage. Impaired brain regions communicate with the SVZ through various chemotactic signals that promote cell migration and differentiation, primarily involving neural stem cells (NSC), neuroblasts, or glioblasts. However, the activation of gliogenesis and the role of newly formed astrocytes in the post-ischemic scenario remain subjects of debate. We have previously demonstrated that adenosine release after brain ischemia prompts the SVZ to generate new astrocytes. Here, we used transient brain ischemia in mice to identify the cellular origin of these astrocytes within the SVZ neurogenic niche and to investigate their role in the pathological process. By combining immunofluorescence, BrdU-tracing, and genetic cell labeling, we tracked the migration of newborn astrocytes, positive for the proteoglycan marker Thbs4, from the dorsal and medial SVZ to the perilesional barrier surrounding the ischemic core, known as the “glial scar”. We found that these Thbs4-positive astrocytes modulate the dense extracellular matrix at the lesion border by both synthesizing and degrading hyaluronan. We also show that while the accumulation of this polymer at the lesion site is sufficient to recruit newborn astrocytes, its degradation at the SVZ correlates with gliogenesis. These findings suggest that newborn astrocytes could be a promising pharmacological target for modulating the glial scar after brain ischemia and facilitate tissue regeneration.

Revised version: Activation of the subventricular zone (SVZ) following cerebral ischemia is one of the brain’s early responses to counteract neuron loss and minimize tissue damage. Impaired brain regions communicate with the SVZ through various chemotactic signals that promote cell migration and differentiation, primarily involving neural stem cells (NSC), neuroblasts, or glioblasts. However, the activation of gliogenesis and the role of newly formed astrocytes in the post-ischemic scenario remain subjects of debate. We have previously demonstrated that adenosine release after brain ischemia prompts the SVZ to generate new astrocytes. Here, we used transient brain ischemia in mice to identify the cellular origin of these astrocytes within the SVZ neurogenic niche and to investigate their role in the pathological process. By combining immunofluorescence, BrdU-tracing, and genetic cell labeling, we tracked the migration of newborn astrocytes, positive for the proteoglycan marker Thbs4, from the dorsal and medial SVZ to the perilesional barrier surrounding the ischemic core, known as the “glial scar”. We found that these Thbs4-positive astrocytes modulate the dense extracellular matrix at the lesion border by both synthesizing and degrading hyaluronan. We also show that while the accumulation of hyaluronan at the lesion site is sufficient to recruit newborn astrocytes, its degradation at the SVZ correlates with gliogenesis. These findings suggest that newborn astrocytes could be a promising pharmacological target for modulating the glial scar after brain ischemia and facilitate tissue regeneration.

Introduction

Earlier version: Cell regeneration in the adult mammalian brain primarily occurs mainly in two neurogenic niches: the subventricular zone (SVZ) in the lateral ventricle and the subgranular zone (SGZ) in the dentate gyrus of the hippocampus (Bond et al., 2015). Adult brain stem cells are multipotent and originate from the embryonic radial glia, thus sharing key cell markers and features with resident astrocytes. In the adult mouse brain, neuroblasts from SVZ migrate through the rostral migratory stream (RMS) to the olfactory bulb (OB), where neurogenesis takes place (Luskin, 1993). Neuronal hyperactivity (e.g., during epilepsy), as well as pathological conditions (e.g., stroke) or natural (e.g., aging) processes, can impair the physiological functions of the neurogenic niches, altering their neurogenic and gliogenic potential (Chaker et al., 2016; Kralic et al., 2005; Sundholm-Peters et al., 2005). For instance, it is known that SVZ cells are activated in vivo after focal ischemia following a middle cerebral artery occlusion (MCAO) (Ceanga et al., 2021).

Revised version: Cell regeneration in the adult mammalian brain primarily occurs in two neurogenic niches: the subventricular zone (SVZ) in the lateral ventricle and the subgranular zone (SGZ) in the dentate gyrus of the hippocampus (Bond et al., 2015). Adult brain stem cells are multipotent and originate from the embryonic radial glia, thus sharing key cell markers and features with resident astrocytes. In the adult mouse brain, neuroblasts from SVZ migrate through the rostral migratory stream (RMS) to the olfactory bulb (OB), where neurogenesis takes place (Luskin, 1993). Neuronal hyperactivity (e.g., during epilepsy), as well as pathological conditions (e.g., stroke) or natural (e.g., aging) processes, can impair the physiological functions of the neurogenic niches, altering their neurogenic and gliogenic potential (Chaker et al., 2016; Kralic et al., 2005; Sundholm-Peters et al., 2005). For instance, it is known that SVZ cells are activated in vivo after focal ischemia following a middle cerebral artery occlusion (MCAO) (Ceanga et al., 2021).

We have previously demonstrated in oxygen and glucose-deprived (OGD) organotypic brain slices that the ischemic region and the SVZ establish a biochemical interplay based on a gradient of toxic and protective factors released early after the insult (Cavaliere et al., 2006). Furthermore, others have demonstrated that brain stroke models in mice can activate astrogliogenesis at the expense of neurogenesis (Benner et al., 2013; Laug et al., 2019; Li et al., 2010). Benner and colleagues identified a SVZ-derived astrocyte population characterized by high levels of the membrane-bound proteoglycan thrombospondin 4 (Thbs4), a cell-cell and cell-matrix adhesion molecule that has been considered neuroprotective against ischemic damage (Benner et al., 2013; Laug et al., 2019). Ablation of Thbs4 in KO mice also produced abnormal ischemic scar formation and increased microvascular hemorrhage, suggesting a role of SVZ astrogliogenesis in brain damage and glial scar formation induced by ischemic stroke (Benner et al., 2013).

The perilesional barrier known as the glial scar assembles early after brain ischemia, isolating the damaged area and limiting both the spread of pathogens and potentially toxic damage-associated molecular patterns (DAMPs) into the surrounding area (Anderson et al., 2016; Bush et al., 1999; Faulkner et al., 2004). Consisting of packed reactive astrocytes and a dense extracellular matrix composed mostly of hyaluronic acid (HA), the glial scar represents a physical barrier for the SVZ-generated newborn neurons that migrate into the ischemic core (Arvidsson et al., 2002; Yamashita et al., 2006). Thus, modulating glial scar formation after brain ischemia may be a strategy to reduce tissue damage and restore brain functionality. In ischemic stroke, the brain extracellular matrix (ECM) goes through a significant reorganization due to injury and must undergo successful remodeling to promote brain repair. ECM elements involved in the ischemic cascade induce activation of astrocytes and microglia, which modulate ECM structure by releasing glycoproteins into the extracellular space (Adams and Gallo, 2018; Fawcett, 2015). This crosstalk culminates in the formation of the glial scar, which is created locally by reactive astrocytes and astrocytes-secreted ECM molecules such as HA and chondroitin sulphate proteoglycans (Bush et al., 1999; Lau et al., 2013).

We have previously demonstrated that high extracellular adenosine levels, a hallmark of ischemic insults, induce astrogliogenesis from the SVZ (Benito-Muñoz et al., 2016). Following our previous results, here we characterize the SVZ response to brain ischemia and investigate the role of newly generated astrocytes in the modulation of the glial scar after MCAO. We show that ischemic conditions increase the number of newly generated astrocytes in the SVZ and their migration to the ischemic area. We also propose the role of these astrocytes in modulating the extracellular matrix at the glial scar and suggest this astrocyte population as a possible target for brain repair after brain ischemia.

Results

Thbs4 is expressed in the neurogenic niche of the SVZ

As suggested by others (Benner et al., 2013; Girard et al., 2014; Pous et al., 2020), Thbs4 is highly expressed in astrocytes of the SVZ and RMS and might play a role in ischemia-induced neurogenesis. To confirm these observations, we labeled the mouse SVZ by immunofluorescence for Thbs4. Thbs4 was mainly expressed in the SVZ, RMS, and OB but also in the cerebellum and ventral tegmental area (Fig. 1a). In the SVZ, the staining was associated with cells with different morphologies (Fig. 1b), ranging from cells with small primary apical processes in contact with the ventricle, to rounded NSC-like cells in the dorsolateral SVZ horn. After quantification by direct count of Thbs4-positive cells, Thbs4-positive astrocytes represented 21.4% of the total SVZ cells (Fig. 1c), with more than 60% of them displaying a radial glia-like type-B morphology (Bond et al., 2015; Doetsch et al., 1999) (Fig. 1d). This is in accordance with other studies reporting Thbs4 expression in B-cells, either by immunohistochemistry (Beckervordersandforth et al., 2010; Benner et al., 2013) or single-cell RNA sequencing (Basak et al., 2018; Cebrian-Silla et al., 2021; Llorens-Bobadilla et al., 2015).

– Thbs4 labels type-B cells in the SVZ neurogenic niche

(a) Schematic drawing of the brain showing Thbs4 labeling throughout the mouse brain. Thbs4 was primarily found in the SVZ and RMS of the adult mouse telencephalon. Arrows indicate non-specific Thbs4 staining in cerebellum and ventral tegmental area. (b) Thbs4 positive cells in the SVZ displayed various morphologies. (c) Proportion of Thbs4-positive cells relative to the total NSCs population in the SVZ. (d) Classification of Thbs4-positive cells in the SVZ based to their NSC morphology. (e) Electroporation of pThbs4-GFP at P1 revealed Thbs4 expression in the postnatal dorsal SVZ. (f) Quantification of Thbs4-reporter expression showing the proportion of NSC expressing Thbs4 postnatally in the dorsal SVZ. (g) Thbs4-GFP reporter expression was upregulated in the rostral SVZ. BV: Blood vessel; LV: Lateral ventricle. Scale bar = 10 µm (b) and 50 µm (e). Bars represent mean ± SEM. *p < 0.05 and **p < 0.01, Tukey post hoc test (after one-way ANOVA was significant at p < 0.05).

Earlier version: To validate the immunohistochemistry protocol for Thbs4, we labeled the SVZ cells with green fluorescent protein (GFP) under the control of the Thbs4 promoter (pThbs4-eGFP) by postnatal (P1) electroporation (Fig. 1E). Twenty-four hours after electroporation, almost 20% of total cells in the dorsal SVZ expressed GFP (Fig. 1F), with a greater expression in the rostral SVZ (Fig. 1G), confirming that Thbs4-positive cells were mainly localized in the neurogenic niche of the SVZ. This also suggests that Thbs4 is expressed in NSCs at early postnatal stages, generating astrocytes alongside development.

Revised version: To validate the immunohistochemistry protocol for Thbs4, we labeled the SVZ cells with green fluorescent protein (GFP) under the control of the Thbs4 promoter (pThbs4-eGFP) by postnatal (P1) electroporation (Fig. 1e). Twenty-four hours after electroporation, almost 20% of total cells in the dorsal SVZ expressed GFP (Fig. 1f), with a greater expression in the rostral SVZ (Fig. 1g), confirming that Thbs4-positive cells were mainly localized in the neurogenic niche of the SVZ. This also suggests that Thbs4 is expressed in NSCs at early postnatal stages, generating astrocytes during development.

Thbs4 expression in the SVZ increases after brain ischemia

Earlier version: To investigate whether brain ischemia could activate astrogliogenesis in the SVZ, we analyzed cell proliferation and Thbs4 expression in the SVZ after MCAO. Sixty minutes of MCAO (Fig. S1A) caused neuronal damage in the cortex and striatum, as observed by cresyl violet staining (Fig. S1B) and NeuN immunofluorescence (Fig. S1C). The survival rate at 28 days post-lesion was around 50% (Fig. S1D). Moreover, motor deficits and animal weight of ischemic mice were monitored every day after MCAO for 30 days. Motor deficits improved overtime, while body weight dropped sharply the first 3 days after MCAO and was partially restored at 30 dpi (Fig. S1E-F).

Revised version: To investigate whether brain ischemia could activate astrogliogenesis in the SVZ, we analyzed cell proliferation and Thbs4 expression in the SVZ after MCAO. Sixty minutes of MCAO (Fig. S1a) caused neuronal damage in the cortex and striatum, as observed by cresyl violet staining (Fig. S1b) and NeuN immunofluorescence (Fig. S1c). The survival rate at 28 days post-lesion was around 50% (Fig. S1d). Moreover, motor deficits and animal weight of ischemic mice were monitored every day after MCAO for 30 days. Motor deficits improved overtime, while body weight dropped sharply the first three days after MCAO and was partially restored at 30 dpi (Fig. S1e, f).

Earlier version: To study cell proliferation in the SVZ, we administered 50 mg/kg BrdU i.p. every two hours (3 doses in total) the day before MCAO (Fig. S2A). We observed an increase in the number of total SVZ BrdU-positive cells 24h after MCAO (Fig. S2B-C). Ki67 immunofluorescence of SVZ wholemount preparations (Fig. S2D) also revealed a more than tenfold increase of proliferative cells 24h after MCAO (Fig. S2E). In addition, we observed a significant increase in cleaved Caspase3-positive cells 24h post-lesion (Fig. S2F-G). These results suggest a fast and specific proliferative response in the ischemia-induced SVZ.

Revised version: To study cell proliferation in the SVZ, we administered 50 mg/kg BrdU i.p. every two hours (three doses in total) the day before MCAO (Fig. S2a). We observed an increase in the number of total SVZ BrdU-positive cells 24h after MCAO (Fig. S2b, c). Ki67 immunofluorescence of SVZ wholemount preparations (Fig. S2d) also revealed a more than tenfold increase of proliferative cells 24h after MCAO (Fig. S2e). In addition, we observed a significant increase in cleaved Caspase3-positive cells 24h post-lesion (Fig. S2f, g). These results suggest a fast and specific proliferative response in the ischemia-induced SVZ.

Earlier version: We next analyzed Thbs4 staining in combination with Nestin, an NSC marker, and Doublecortin (DCX), a marker of neuroblasts, before and after MCAO. SVZ was analyzed in sham and ischemic mice 7-, 15- and 30-days post-ischemia (dpi). The immunofluorescence analysis of the SVZ (Fig. 2A) showed a transient decrease in Nestin and DCX-positive cells but an increase of Thbs4-positive astrocytes up to 15 dpi (Fig. 2B-D) compared to sham animals. The increase of Thbs4 astrocytes was also confirmed by Western blot analysis using SVZ homogenates (Fig. 2E-F), evidencing a 2-fold increased expression of Thbs4 in the ischemic SVZ. Similarly to immunofluorescence, Western blot also showed decreased Nestin protein levels at 15 dpi (Fig. 2G). These results demonstrate that Thbs4 expression in the SVZ increases after ischemia, coinciding with a decrease in neuroblast markers, suggesting that the classical pro-neuroblast program of the SVZ is interrupted in favor of the generation of astrocytes after ischemia.

Revised version: We next analyzed Thbs4 staining in combination with Nestin, an NSC marker, and Doublecortin (DCX), a marker of neuroblasts, before and after MCAO. The entire SVZ was analyzed in sham and ischemic mice 7-, 15- and 30-days post-ischemia (dpi). The immunofluorescence analysis of the SVZ (Fig. 2a) showed a transient decrease in Nestin and DCX-positive cells but an increase of Thbs4-positive astrocytes up to 15 dpi (Fig. 2b-d) compared to sham animals. The increase of Thbs4 astrocytes was also confirmed by Western blot analysis using SVZ homogenates (Fig. 2e, f), evidencing a 2-fold increased expression of Thbs4 in the ischemic SVZ. Similarly to immunofluorescence, Western blot also showed decreased Nestin protein levels at 15 dpi (Fig. 2g). These results demonstrate that Thbs4 expression in the SVZ increases after ischemia, coinciding with a decrease in neuroblast markers, suggesting that the classical pro-neuroblast program of the SVZ is interrupted in favor of the generation of astrocytes after ischemia.

– Thbs4 expression is upregulated in the SVZ after brain ischemia

(a) Representative images of Thbs4 (green), Nestin (red) and DCX (magenta) in the SVZ under physiological conditions (left) and 15 days post-injury (dpi) (right). (b, c) The number of actively proliferating NSCs and neuroblast decreases at 7 and 15 dpi in the SVZ, as shown by reduced Nestin (b) and DCX (c) levels, respectively. (d) Thbs4 expression increases in the SVZ over time following ischemic injury. (e) Representative western blot images of Nestin and Thbs4 in sham and 15 dpi SVZ. (f, g) Quantification of mean gray value (MGV) from (e), normalized to GAPDH, shows an increased in Thbs4 (f) and a decrease in Nestin (g) in the 15 dpi SVZ. Scale bar = 100 µm in (a). n = 6 (b, c, d) and 3 (e, f, g) per condition. Bars represent mean ± SEM. *p < 0.05 and **p < 0.01; two-tail Student’s t test (sham vs. 15ddMCAO) and Tukey post-hoc test (after one-way ANOVA was significant at p < 0.05). See also Supplementary Figures 1 and 2.

Characterization of ischemia-induced cell populations in the SVZ

To characterize the different cell populations of the SVZ activated after focal ischemia, we labeled the whole pool of proliferating cells with 1% BrdU in drinking water for 14 days before MCAO (Codega et al., 2014). One month after MCAO, BrdU in the SVZ primarily stained slow proliferative cells like type B cells, whereas in rapid transit-amplifying cells (type C), BrdU was diluted away. Therefore, we chose to label type C cells by intraperitoneal administration (50 mg/kg) of 5-Iodo-2′-deoxyuridine (IdU) 24h before euthanasia (Fig. 3a). To identify the cellular origin of newborn astrocytes, we analyzed the SVZ 30 dpi by immunofluorescence co-labeling Thbs4 and GFAP together with BrdU or IdU. We found a 2-fold increase in the total amount of proliferative cells (BrdU-positive) after MCAO with respect to the sham group (Fig. 3b), whereas we did not observe a significant change in the number of IdU-positive cells (Fig. 3c). The number of slow proliferative NSC (Thbs4/GFAP/BrdU) increased linearly after MCAO (Fig. 3d, e) especially in the dorsal region of the SVZ (Fig. S3a), whereas the number of activated NSC (Thbs4/GFAP/IdU) did not change significantly, either in the whole SVZ (Fig. 3f) or in any subregion analyzed (Fig. S3b). These results suggest that, after ischemia, Thbs4-positive astrocytes derive from the slow proliferative type B cells.

Earlier version: Thbs4-positive cells derive from B-type NSCs in the SVZ after brain ischemia.

Revised version: – Thbs4-positive cells derive from type-B NSCs in the SVZ after brain ischemia

(a) Experimental design: chronic BrdU (1% in water) treatment was administered for two weeks. BrdU labeling persisted only in B-type NSCs 30 days after treatment. Animals were injected with IdU (50 mg/kg) three times the day before euthanasia to label proliferating cells. (b, c) BrdU (b) and IdU (c) positive cells in the SVZ. Only BrdU-positive cells showed a significant increase at 30 dpi. (d) Representative confocal images showing Thbs4, GFAP and BrdU levels in sham (top) and 30 dpi mice (bottom). (e, f) Thbs4/GFAP/BrdU positive cells (slow proliferative cells) increased in the SVZ at 30 dpi (e), while no changes were observed in Thbs4/GFAP/IdU positive cells (fast proliferative cells). Scale bar = 10 µm (d). n = 4 (sham) and 3 (MCAO). Bars represent mean ± SEM. *p < 0.05 by two-tail Student’s t-test (sham vs. 30dd MCAO). See also Supplementary Figure 3.

Ischemia-induced newborn astrocytes migrate from the dorsal SVZ to ischemic regions

To confirm that Thbs4-positive astrocytes originate from type B cells of dorsal SVZ, we labeled NSCs of dorsal SVZ by electroporating the pCAGGSx-CRE plasmid in P1 Ai14 Rosa26-CAG-tdTomato transgenic mice (Fig. 4a). Plasmid electroporation induced the expression of tdTomato (tdTOM) by Cre recombination, enabling the spatial tracking and cell fate assessment of electroporated NCSs from the dorsal SVZ. Three to four months after electroporation, we performed the MCAO, and tdTOM fluorescence was analyzed at 7, 30, and 60 dpi. To confirm the efficient labeling of NSCs by the electroporation protocol, we followed the tdTOM fluorescence in the OB (Fig. S4a, see methods). Only those animals with tdTOM-positive cells in the OB were taken into account for later analysis. After quantification of tdTOM-positive NSCs 7, 30, and 60 dpi, we observed a significant decrease in tdTOM-positive cells in ischemic animals in all SVZ regions analyzed (Fig. 4b, c and Fig S4b), suggesting that they either left the area, or died. Conversely, there was a general increase of tdTOM-positive cells outside the SVZ (Fig. 4d, e), especially in the intermediate and caudal axis of the whole brain compared to sham animals (Fig. S4d-g), suggesting that electroporated cells had migrated away from the SVZ.

– Ischemia-induced Thbs4 astrocytes migrate from the SVZ to ischemic areas

(a) Experimental design: Ai14, ROSA26-CAG-tdTomato transgenic mice were electroporated in the dorsal SVZ at P1. The pCAGGSx-CRE plasmid induced tdTomato (tdTOM) expression in dorsal NSCs. MCAO was performed 3-4 months later in the same electroporated hemisphere. Mice were analyzed at 7, 30, and 60 dpi. (b) Quantification of tdTOM-positive cells shows a decrease in the dorsal SVZ following brain ischemia. (c) Representative images of tdTOM-positive cells in sham (top) and 60 dpi SVZ (bottom). (d) Quantification of tdTOM-positive cells in the whole brain shows a gradual increase outside the SVZ after brain ischemia. (e) Representative images of tdTOM expression in sham (left) and 60 dpi mice (right). (f) Representative image of Thbs4/tdTOM-positive cells in the dorsal SVZ. (g) Quantification shows an increase in Thbs4/tdTOM-positive cells in the SVZ at 30 and 60 dpi. (h) The proportion of Thbs4-positive cells within the tdTOM-positive pool increases in the SVZ at 30 and 60 dpi. (i) Around 50% of tdTOM-positive cells expressed Thbs4 in the infarcted tissue at 60 dpi. (j) Representative image of Thbs4 and tdTOM expression near the damaged area. Scale bar = 50 µm (c), 10 µm (f), and 100 µm (g). n = 6 animals per condition. Bars represent mean ± SEM. *p < 0.05, **p < 0.01, and ***p < 0.001 by Chi-square test (h) and Tukey post-hoc test (after one-way ANOVA was significant at p < 0.05). See also Supplementary Figure 4.

To identify the proportion of NSCs that differentiated into Thbs4-positive astrocytes following ischemia, we analyzed the co-expression of Thbs4 with tdTOM after MCAO. Despite the general decrease of tdTOM-positive cells in the SVZ, the fraction of Thbs4/tdTOM cells increased at 30 dpi after MCAO (Fig. 4f, g), reaching a maximum increase at 60 dpi (Fig. 4h), in particular in rostral and intermediate SVZ (Fig. S4c). Finally, we analyzed the Thbs4-positive astrocytes recruitment to the ischemic regions. More than half of the total tdTOM-positive cells in the infarcted area expressed the Thbs4 marker at 60 dpi (Fig. 4i, j). These results suggest that ischemia-induced astrogliogenesis in the SVZ occurs in type B cells from the dorsal region and that these newborn Thbs4-positive astrocytes migrate to the ischemic areas.

Thbs4-positive astrocytes accumulate at the glial scar after brain ischemia

Besides its role in the neurogenic niches, Nestin is also rapidly and locally expressed by reactive astrocytes and participates in glial scar formation (Frisén et al., 1995). Thus, we used Nestin immunofluorescence as a reliable tool to identify the lesion border where the glial scar occurs (Fig. 5a). Nestin staining of the scar was also compared with more classical methodologies like TTC staining (Fig. 5b), where viable tissue is stained red and dead tissue is left unstained. The white-pale pink area was then measured as infarct tissue by TTC staining (Fig. 5c). Unlike the TTC staining, which can only be performed shortly after the lesion, the visualization of the glial scar with Nestin immunofluorescence allowed us to analyze the lesion at later time-points. We observed a rapid formation of the glial scar as early as 7 dpi (Fig. 5d). Analyzing the expression of the two astrocytic markers, Thbs4 and GFAP, as a percentage of the glial scar region-of-interest (ROI), we observed different dynamics and localization. GFAP expression increased rapidly after 7 days within the glial scar ROI, whereas Thbs4 was upregulated at 15 dpi (Fig. 5e). This delay suggested that the glial scar harbors two different astrocytic sub-populations: GFAP-only or GFAP/Thbs4. To analyze the spatial location of GFAP and Thbs4 markers within the Nestin border, we measured the fluorescence intensity profile of damaged area at 30 dpi. GFAP was found in both the scar and core areas, whereas Thbs4 fluorescence was found mostly at the borders of the glial scar (Fig. 5f, g). Analysis and dynamics of Thbs4-positive astrocyte migration were further performed in the corpus callosum (CC), striatum (Str), cortex (Ctx) and olfactory bulb (OB) (Fig. S5a). We observed a linear increase of Thbs4-positive astrocytes from 7 to 60 dpi, both in the Ctx and Str (Fig. S5b). Migration occurred mainly through the corpus callosum; among the total GFAP-positive astrocytes, 36.5% were also positive for Thbs4, suggesting that more than one-third of the astrocytes in the glial scar originated from the SVZ (Fig. S5c). Likewise, the number of DCX-positive cells dramatically decreased 15 dpi in all OB layer (Fig. S5d-f) together with a significant decrease of BrdU-positive cells (Fig. S5g) DCX/BrdU-positive neuroblasts in the RMS (Fig. S5h), suggesting that ischemia induced a change in the neuroblasts’ ectopic migratory pathway, depriving the OB layers of SVZ newborn neurons.

– Ischemia-induced Thbs4 astrocytes accumulate at the glial scar

(a) The glial scar could be observed at 30 dpi using Nestin immunofluorescence. (b) Classical TTC technique underestimated ischemic regions compared with Nestin immunofluorescence. (c) Representative images of TTC staining in infarcted brains. (d) Representative confocal images showing the time course of Nestin expression at 3, 7, 15 and 30 dpi. (e) The Nestin-positive glial scar was established by 15 dpi, while Thbs4 astrocytes arrived at the glial scar later, by 30 dpi. (f) Fluorescence profile of Thbs4, GFAP and Nestin marker from the core to penumbra areas shown Thbs4 was mainly located at the borders of the Nestin-positive glial scar. (g) Representative images of Thbs4 (green), GFAP (white) and Nestin (red) markers in sham (top) and 30 dpi mice (bottom). Scale bar = 50 µm (g). n = 3 animals per condition. Bars/lines represent mean ± SEM. *p < 0.05 and **p < 0.01 by two-tail Student’s t test (TTC vs. Nestin) and by Dunnett posthoc test (after two-way ANOVA was significant at p < 0.05). See also Supplementary Figure 5.

Thbs4-positive astrocytes modulate extracellular matrix homeostasis at the glial scar

To investigate the role of Thbs4-positive astrocytes in the glial scar, we analyzed hyaluronic acid (HA), the main component of the brain extracellular matrix (ECM), using the hyaluronic acid binding protein (HABP) as a marker. Following ischemic damage, local astrocytes (GFAP+/Thbs4-) react by adopting a fibrotic morphology, producing high-molecular weight HA (HMW-HA) to form the glial scar and isolate the damaged area (Preston and Sherman, 2011). We then examined HA deposition in the corpus callosum and the infarcted areas after MCAO. HABP immunofluorescence (Fig. 6a) and image analysis revealed HA accumulation in the infarcted areas. Skeleton analysis, which quantifies the length of HA chains and correlates with its molecular weight, identified progressively longer HA cable-like structures at the glial scar as the infarct evolved (Fig. 6b). Fractal dimension analysis, which reflects the interconnectivity of the extracellular matrix (Soria et al., 2020), further confirmed HA accumulation and increased matrix complexity in infarcted areas, but not in corpus callosum (Fig. 6c), suggesting that ECM accumulation is different in white and gray matter.

– Thbs4-positive astrocytes produce hyaluronan at the glial scar after MCAO

(a) Representative images of hyaluronic acid (HA) labeled with HABP in the contralateral (1) and ipsilateral hemisphere (2). Only perineuronal nets are clearly visible in the healthy tissue, while a dense interstitial matrix is observed in ischemic areas. (b) Cumulative distribution of skeletonized HA signal in sham and infarcted regions of 7, 15 and 30 dpi mice (p < 0.0001). (c) Quantification of fractal dimension showed increased extracellular matrix interconnectivity in the infarcted cortex and striatum (Ctx & Str) at 30 dpi. (d) Representative confocal images of Thbs4 (green) and HABP (red) in the cortex and striatum of sham and 30 dpi mice. Note the HA labelling on the surface of Thbs4-positive cells. (e) Fluorescence intensity profile of Thbs4, GFAP and HABP markers from the lesion core to penumbra areas. Thbs4 expression was mainly restricted to the borders of the HA-positive glial scar at 30 dpi. (f) Representative confocal images of Thbs4 and HABP markers in sham and 30 dpi mice (top). Examples of HA (red) occupying the membrane mask of Thbs4-positive astrocytes in sham and 30 dpi mice (bottom), showing increased HA synthesis after MCAO. (g) Quantification of membrane-bound HA showed increased HA along Thbs4-positive membrane after brain ischemia (corpus callosum: pink; infarcted cortex: blue; infarcted striatum: green). Scale bar = 10 µm (e). n = 6 animals per condition. *p < 0.05 and ****p < 0.0001 by Krustal Wallis test and two-way ANOVA (significant at p < 0.05). See also Supplementary Figures 6 and 7.

The proximity of Thbs4-positive astrocytes to HA in the ischemic scar (Fig. 6d, e) led us to investigate if this population contributed to ECM production and, consequently, to scar formation. Since HA is produced exclusively at the cell membrane (Sherman et al., 2015; Zimmermann and Dours-Zimmermann, 2008), we analyzed HA labeling within 1 µm of the Thbs4-positive cell membrane in high-resolution confocal images (Fig. 6f). Image analysis showed a significant increase in HABP area at Thbs4-positive astrocyte-cell membranes after MCAO (Fig. 6g), suggesting a role of Thbs4-positive astrocytes in HA production within the glial scar.

To explore the potential role of Thbs4-positive astrocytes in HA removal, we quantified intracellular HABP spots in Thbs4-positive astrocytes at the glial scar. While HMW-HA cleavage occurs extracellularly, HA fragments can be internalized for degradation by lysosomal hyaluronidases (Preston and Sherman, 2011; Zimmermann and Dours-Zimmermann, 2008). Intracellular HABP spots were compared between Thbs4-positive astrocytes and local astrocytes (GFAP-positive/Thbs4-negative) (Fig. S6a). Thirty days after MCAO, Thbs4-positive astrocytes showed more intracellular HABP than local astrocytes in the infarcted cortex (Fig. S6b). However, the number of HA spots internalized was higher in local astrocytes compared with Thbs4 astrocytes (Fig. S6c), suggesting that Thbs4-positive cells internalize HA in fewer vesicles. To confirm that astrocytes were degrading HA, we performed an immunofluorescence analysis of CD44 (Fig. S6d). CD44 is a transmembrane glycoprotein receptor that mediates HA internalization and degradation, and therefore ECM remodeling (Dzwonek and Wilczynski, 2015). We detected that around 50% of the total Thbs4-positive astrocytes contained internalized CD44 receptor, in the striatum 30 dpi (Fig. S6e), with higher numbers (60%) for the infarcted corpus callosum. These findings suggest that Thbs4-positive astrocytes participate in both the formation and degradation of the extracellular matrix at the glial scar.

Earlier version: To further investigate HA internalization by astrocytes, we conducted in vitro experiments using rat SVZ neural stem cells cultures. An oxygen and glucose deprivation (OGD) protocol was applied as an ischemia model and cells were fixed 7 days after OGD (Fig. S7A). HA was analyzed via HABP immunofluorescence (Fig. S7B), both inside and outside the Thbs4-positive cells (Fig. S7C). Thbs4-positive cells showed increased intracellular HA after OGD (Fig. S7D), while extracellular HA levels decreased (Fig. S7E). However, this phenomenon was reversed after adding hyaluronidases in the extracellular space (Fig. S7F-G), suggesting that Thbs4-positive cells from SVZ homogenates can internalize or synthesize HA depending on external stimuli. HA internalization was also measured in Iba1-positive cell cultures where we observed a reduced internalization compare to Thbs4-positive cells (Fig. S7H). Lastly, metalloproteinase activity was found to be increased in Thbs4 cultures following OGD (Fig. S7I). Altogether, these results suggest Thbs4-positive cells can internalize HA after ischemic conditions but also synthetize it when ECM is degraded.

Revised version: To further investigate HA internalization by astrocytes, we conducted in vitro experiments using cultures of rat SVZ neural stem cells, a more versatile preparation that displays the same HA internalization and degradation machinery as mouse brain tissue (Sherman et al., 2015). An oxygen and glucose deprivation (OGD) protocol was applied as an ischemia model and cells were fixed 7 days after OGD (Fig. S7a). HA was analyzed via HABP immunofluorescence (Fig. S7b), both inside and outside the Thbs4-positive cells (Fig. S7c). Thbs4-positive cells showed increased intracellular HA after OGD (Fig. S7d), while extracellular HA levels decreased (Fig. S7e). However, this phenomenon was reversed after adding hyaluronidases in the extracellular space (Fig. S7f, g), suggesting that Thbs4-positive cells from SVZ homogenates can internalize or synthesize HA depending on external stimuli. HA internalization was also measured in Iba1-positive cell cultures where we observed a reduced internalization compare to Thbs4-positive cells (Fig. S7h). Lastly, metalloproteinase activity was found to be increased in Thbs4 cultures following OGD (Fig. S7i). Altogether, these results suggest Thbs4-positive cells can internalize HA after ischemic conditions but also synthetize it when ECM is degraded.

Hyaluronan accumulation is sufficient but not exclusive for the recruitment of Thbs4-positive astrocytes

Since Thbs4-positive astrocytes migrate to the glial scar but not to the ischemic core, we hypothesized that HA, the primary component of the ischemic scar, may be the trigger signal for recruiting Thbs4-positive astrocytes. To test this hypothesis, we simulated the ECM accumulation that occurs at the scar by overexpressing the Hyaluronan synthase 2 (HAS2) from the naked mole rat (Heterocephalus glaber), a tool often used in mice to produces HA of very high molecular weight (Tian et al., 2013; Zhang et al., 2023). We used a viral vector (AAV2/9-CAG-HAS2-3xFlag) with a FLAG tag to visualize transduced cells after intracerebral inoculation in the striatum and checked HA accumulation in FLAG-positive areas using HABP staining (Fig. 7a). We overexpressed HAS2 (or GFP as control) unilaterally. In a subset of animals that underwent MCAO, AAV-HAS2 was injected contralateral to the infarct, to compare Thbs4-positive astrocytes recruitment by both stimuli, the inflammation from the ischemic damage and the HA accumulation per se (Fig. 7b). FLAG-positive cells were quantified as a surrogate marker of viral infection for both control and MCAO groups (Fig. 7c). The GFP control group did not increase HA production per se (Fig. 7d). However, HA overexpression was observed in HAS2 groups (both control and MCAO group) by HABP immunofluorescence (Fig. 7e, f).

– Hyaluronan accumulation is sufficient but not exclusive to recruit Thbs4 astrocytes

(a) Representative images of HA accumulation caused by viral-mediated overexpression of HAS2 from naked mole rat. (b) Experimental design: HAS2 viral injections were performed 2 weeks before MCAO. Animals were sacrificed 30 days after MCAO. We used standard GFP virus as control. Two groups of HAS2-injected mice were studied, with only one subjected to MCAO. (c) FLAG-positive cells (HAS2 virus tag) did not show changes in either HAS2 control (HAS2Control) or HAS2-MCAO group (contralateral hemisphere). (d, e, f) Representative images of HABP, GFP and Thbs4 markers in GFP control virus (d), HAS2 control virus (e) and HAS2 MCAO group (f). (g) Thbs4 expression increased after HAS2 viral-mediated overexpression. However, Thbs4 was upregulated in the ischemic hemisphere when MCAO was performed together with the HAS2 viral infection (MCAO ipsi). GFP control virus did not induce increase of Thbs4. (h) HABP was used as a control for HA accumulation. Thbs4 astrocytes in HAS2 MCAO group increased in the ischemic hemisphere even though HA accumulation in both hemispheres (HAS2 and MCAO hemisphere), suggesting additional ischemia-induced signals recruit Thbs4 astrocytes to infarcted areas. Scale bar = 100 µm (a). n = 6 animals per condition. *p < 0.05, **p < 0.01 and ***p < 0.001 by two-tail Student’s t-test (HAS2Control vs. MCAO Contra) and Tukey posthoc test (after one-way ANOVA was significant at p < 0.05).

HA accumulation was sufficient to recruit the Thbs4-positive astrocytes from the ipsilateral SVZ 45 days after AAV-HAS2 inoculation (Fig. 7g), with a similar time window to that observed after MCAO. Interestingly, when we induced a lesion by MCAO and performed viral overexpression of HAS2 in the contralateral hemisphere, we found that Thbs4-positive astrocytes preferentially migrated to the ischemic glial scar (Fig. 7g). This migration of Thbs4-positive astrocytes to the site of ischemic damage occurred despite no differences in HA overexpression between contralateral (HAS2 infected) and ipsilateral (infarcted area) hemispheres (Fig. 7h). These results suggest that the accumulation of HMW-HA alone is sufficient to activate and recruit Thbs4-positive astrocytes from the SVZ. However, the preferential migration to the ischemic scar indicates that brain ischemia generates a more potent biochemical signal that not only activates the SVZ but also recruits newborn astrocytes.

Ischemia-induced gliogenesis at the SVZ correlates with local hyaluronan degradation

It is known that the niche microenvironment has a unique ECM composition (Kjell et al., 2020) and that the stiffness of the niche plays a crucial role in the proliferation and differentiation of progenitor cells (Long and Huttner, 2019; Segel et al., 2019). After observing the close association between Thbs4-positive astrocytes and the ECM at the glial scar, we decided to study the Thbs4/HA tandem in the SVZ following MCAO (Fig. 8a). We did not observe any significant changes in the HA area across the entire SVZ but noted a decrease at 7 and 15 dpi in the dorsal SVZ (Fig. 8b-d). Skeleton analysis showed fragmentation of HA exclusively in the dorsal SVZ at 7 and 15 dpi (Fig. 8e-g), i.e., earlier than the migration of these newborn astrocytes toward the lesion.

– HA degradation in the SVZ correlates with Thbs4 response to brain ischemia

(a) Representative image of Thbs4 and HA (HABP) markers in the SVZ at 30 dpi. (b) Representative images of Thbs4 and HA in the sham and 15, 30 dpi SVZ. (c, d) HA levels did not change in the entire SVZ after MCAO (c) but decreased significantly in the dorsal SVZ (d). Thbs4 expression increased throughout the SVZ after MCAO (c), with a marked rise at 30 dpi in the dorsal SVZ (d), coinciding with HABP reduction. (e) Representative confocal (top) and skeletonized (bottom) images of HABP in the dorsal SVZ of sham, 15 and 30 dpi mice. (f) Cumulative distribution of HA skeletons showed a reduction in the length of HA cable-like structures only in the dorsal SVZ at 7 and 15 dpi. (g) The interconnectivity of the extracellular matrix decreased by 7 dpi in the dorsal SVZ, as measured by fractal dimension. (h) Experimental design for qPCR: 3–4-month-old mice were subjected to MCAO, and fresh SVZ tissue was extracted 15 and 30 days after MCAO for RT-qPCR analysis. (i) Hyaluronan degradation genes (Hyal1, Hyal2 and Hyal3) increased by 15 and 30 dpi in the SVZ, while hyaluronan synthase genes (HAS1 and HAS2) were overexpressed later, at 30 dpi. Thbs4 and Hif1a were also upregulated by 15 and 30 dpi in the SVZ. Scale bar = 100 µm (a). n = 6 (c, d, e and g) and 3 (i). *p < 0.05, **p < 0.01 and ***p < 0.001 by Krustal Wallis test and Tukey post-hoc test (after one-way ANOVA was significant at p < 0.05).

Next, we examined the expression of genes related to HA degradation (Hyase 1, 2 and 3) and synthesis (HAS 1 and HAS 2). We also analyzed Thbs4 and hypoxia-inducible factor 1-alpha (Hif1a), which are known to be upregulated in the SVZ following brain ischemia (Benito-Muñoz et al., 2016; Liu et al., 2007). Fresh SVZ tissue was collected from sham, 15 and 30 dpi mice (Fig. 8h). We observed an increased expression of all Hyases at 15 dpi and slightly decreased by 30 dpi, whereas synthases (HAS) showed a gradual increase from 15 to 30 dpi (Fig. 8i). Additionally, we observed an upregulation of Thbs4 and Hif1a both at 15 and 30 dpi (Fig. 8i). These results suggest that HA degradation in the SVZ is correlated with the Thbs4 response to ischemia.

Discussion

The effects of brain insult on adult neurogenesis have been extensively discussed, both in animal models and humans. However, whether these injuries can trigger the generation of newborn astrocytes from the neurogenic niches remains a topic of debate. Neural stem cells in the adult brain are a remnant of embryonic radial glia. Some researchers suggest that depletion of neurogenesis in the adult brain is a consequence of astrogliogenesis (Encinas et al., 2011), while others propose astrogliogenesis is part of a neurorestorative program originating from the neurogenic niches (Bonzano et al., 2018; Sohn et al., 2015). Despite these differing perspectives, direct and definitive evidence of the generation of newborn astrocytes from the SVZ or the SGZ has yet to be presented, as well as the potential role of these astrocytes in aging and pathological conditions. In this study, we confirm that a strong and acute damage-associated stimulus in the brain, such as an ischemic infarct, prompts astrogliogenesis in the SVZ. Furthermore, we demonstrate that these astrocytes, which are positive for Thbs4, migrate to the glial scar and play a role in its modulation, being capable of synthesis and degradation of hyaluronan, a key component of the ECM at the scar.

Our findings align with previous studies, where cortical injury initiated by photothrombotic/ischemia was shown to trigger a substantial production of Thbs4-positive astrocytes from the SVZ, a process modulated by Notch activation (Benner et al., 2013). In this work, the authors suggested a clear role for Thbs4-positive astrocytes in modulating the ischemic scar. In fact, homozygotic deletion of the Thbs4 gene led to abnormal glial scar formation and increased microvascular hemorrhage following brain ischemia. Similarly, our results demonstrate that the adult neurogenic niche at the SVZ responds early after ischemic stroke, degrading local hyaluronan plausibly to modify niche stiffness and facilitate proliferation and differentiation of NSCs. Further experiments using Thbs4-knockout models will help confirm this hypothesis.

By labeling NSC via in vivo postnatal electroporation, we provided proof of the origin and migration of newborn Thbs4-positive astrocytes from the SVZ to the ischemic scar in response to the injury. Notably, Thbs4 is expressed in the newborn astrocytes that migrate to the scar, whereas local reactive astrocytes at the site of injury express only GFAP. It is likely that SVZ-derived newborn astrocytes require Thbs4 expression to facilitate their migration to the ischemic area, as Thbs4 deletion has been shown to impair migration of SVZ-derived cells along the RMS (Girard et al., 2014).

It is noteworthy that while the local, “GFAP-only”, reactive astrocytes at the scar are capable of forming the required ECM (Cregg et al., 2014; Wanner et al., 2013), we show that Thbs4-positive (Thbs4/GFAP) astrocytes are also recruited from the SVZ to modulate scar formation. Unlike local reactive astrocytes, which are hastily generated and present both at the core and periphery of the lesion, the newborn astrocytes from the SVZ arrive later to the injury. We demonstrate that these Thbs4-positive astrocytes migrate from the SVZ to the lesion and halt at the lesion border (precisely defined by Nestin staining), establishing residence at the scar but not beyond.

Earlier version: The molecular signals that trigger gliogenesis in the SVZ and recruit newborn astrocytes to the lesion remain unclear. We previously identified adenosine, a byproduct of ATP hydrolysis, as a strong candidate, as it stimulates astrogliogenesis through A1 receptors following OGD in vitro and ischemia in vivo (Benito-Muñoz et al., 2016). Here, we show that viral-mediated overproduction of HA, the primary component of the glial scar, is sufficient to activate the generation and migration of the Thbs4-positive astrocytes from the SVZ. However, our findings also show that the ischemic infarct, with its proinflammatory (cytokines, chemokines) and proteolytic (MMPs, ROS) signals (Jayaraj et al., 2019), serves as a more potent attractor than the ECM alone.

Revised version: The molecular signals that trigger gliogenesis in the SVZ and recruit newborn astrocytes to the lesion remain unclear. We previously identified adenosine, a byproduct of ATP hydrolysis, as a strong candidate, as it stimulates astrogliogenesis through A1 receptors following OGD in vitro and ischemia in vivo (Benito-Muñoz et al., 2016). Here, we show that viral-mediated overproduction of HA, the primary component of the glial scar, is sufficient to activate the generation and migration of the Thbs4-positive astrocytes from the SVZ. Loss-of-function studies, using HA-depletion models or HA synthase (Has)-deficient mice are still needed to corroborate this finding, although the inflammation associated with HA deficiency might confound interpretation. Indeed, our results also show that the ischemic infarct, with its proinflammatory (cytokines, chemokines) and proteolytic (MMPs, ROS) signals (Jayaraj et al., 2019), serves as a more potent attractor than the ECM alone.

Earlier version: The glial response from the SVZ we propose in this study can have several functional implications. Thbs4-positive astrocytes generated from the SVZ might act as a second wave of reactivity, supporting local astrocytes in modulating ECM formation. Hyper-reactive astrocytes at the lesion border decrease over time, along with reductions in HA synthesis and CD44 activity, suggesting that the second wave of Thbs4-positive astrocytes from the SVZ could sustain high levels of HA production (Lindwall et al., 2013). Thus, Thbs4-positive astrocytes may be crucial during this delayed temporal window when glial scar remodeling is essential for neurological recovery (Cai et al., 2017; Dzyubenko et al., 2018). However, it’s important to note that only about one-third of the astrocytes at the glial scar are derived from the SVZ. This suggests that while SVZ-derived astrocytes may be important for glial scar maintenance, local astrocytes remain the central players in this process.

Revised version: The glial response from the SVZ we propose in this study might have several functional implications. Thbs4-positive astrocytes generated from the SVZ might act as a second wave of reactivity, supporting local astrocytes in modulating ECM formation. Hyper-reactive astrocytes at the lesion border decrease over time, along with reductions in HA synthesis and CD44 activity, suggesting that the second wave of Thbs4-positive astrocytes from the SVZ could sustain high levels of HA production (Lindwall et al., 2013). Thus, Thbs4-positive astrocytes may be crucial during this delayed temporal window when glial scar remodeling is essential for neurological recovery (Cai et al., 2017; Dzyubenko et al., 2018). However, it’s important to note that only about one-third of the astrocytes at the glial scar are derived from the SVZ. This suggests that while SVZ-derived astrocytes may be important for glial scar maintenance, local astrocytes remain the central players in this process.

Based on our observations of both HA production and internalization by Thbs4-positive astrocytes in vivo within the ischemic region and in vitro following OGD, we propose a role for these astrocytes in the modulation, and possibly maintenance, of the ECM at the scar. It is likely that ECM degradation at the scar is also carried out by professional macrophages, such as microglia, which are known to be recruited to the scar (Fawcett and Asher, 1999) and have been shown to degrade ECM through phagocytosis (Nguyen et al., 2020; Soria et al., 2020). Understanding which cells within the scar environment participate in its modulation, particularly its degradation, holds highly therapeutic value for promoting regeneration (Dzyubenko et al., 2018; Silver and Miller, 2004). Here, we shed light on the unknown role of SVZ-derived astrocytes activated by ischemic brain injury and propose them as critical players in tissue regeneration after brain ischemia.

Added in the revised version: Methods

Added in the revised version: Animals

Added in the revised version: Male and female wild-type C57BL6/J mice and Ai14-Rosa26-CAG-tdTomatoflx/flx transgenic mice bred at the animal facility of UPV/EHU were used for in vivo experiments. Sprague Dawley rat pups from P4 to P7 were used for in vitro experiments. All animals for in vitro and in vivo studies were handled in accordance with the European Communities Council Directive (2010/63/EU). All possible efforts were made to minimize animal suffering and the number of animals used. Animals were maintained under standard laboratory conditions with food and water ad libitum.

Added in the revised version: Transient MCAO

Added in the revised version: Brain ischemia was induced in 3-4-month old C57BL/6 mice (approximately 25 g) by occluding 60 minutes of the lenticulo-striate arteries with a 10 mm length of silicone-coated 6-0 monofilament nylon suture (602223PK10, Doccol Corporation) using a modification of the protocol described by (Gelderblom et al., 2009). Arteries occlusion generated a specific lesion in the cortex and striatum that was evaluated at different times after the MCAO (7, 15, 30 or 60 days after occlusion). Ischemic damage was evaluated in the ipsilateral hemisphere by Cresyl violet staining, Nestin immunofluorescence or hyaluronan staining. Sham animals, in which arteries were visualized but not disturbed, were used as an experimental control. Both experimental groups, sham and ischemic, were anesthetized with a mix of isoflurane and oxygen (induction of anesthesia with 4% isoflurane and maintenance for surgery at 1.5% isoflurane) and maintained under analgesia with buprenorphine 24 hours after ischemia (0.03 mg/Kg body weight i.p. every 12h). Animal wellness was evaluated daily, controlling the body weight and the motor deficit of the animal by means of neurological scores and pole test. Motor deficit was assessed in each animal 1h after MCAO and later every 24h. Animals without neurological deficits were excluded from the study.

Added in the revised version: 5-bromo-2’-deoxyuridine (BrdU) protocol

Added in the revised version: The total pool of proliferating cells was labelled with three injections of 50 mg/kg 5-bromo-2’-deoxyuridine (BrdU) the day before MCAO protocol. Animals were sacrificed the day after MCAO to analyse acute SVZ proliferation. Chronic BrdU protocol was performed for quiescent NSC identification (Codega et al., 2014). Briefly, 1% of BrdU was administered in the drinking water for 14 days. On day 15, BrdU was removed and animals were subjected to MCAO protocol. Mice were sacrificed and perfused at 3, 7, 15, and 30 days after MCAO, although only animals with a temporal window of 30 days between MCAO and perfusion were used for statistical analysis. Activated cells were labelled by i.p. injection of 50 mg/Kg 5-iodo-2′-deoxyuridine (IdU) 24 hours before euthanasia.

Added in the revised version: Electroporation

Added in the revised version: Ai14-Rosa26-CAG-tdTomatoflx/flx transgenic mice and pCAGGS-CRE plasmid were kindly provided for Prof. Harold Cremer and Marie-Catherine Tiveron from Institut de Biologie du Développement de Marseille (IBDM). Postnatal electroporation was performed as described previously (Platel et al., 2019). Mice pups from P0-P2 were anesthetized by hypothermia. pCAGGS-CRE plasmid was used at a final concentration of 2.5 µg/ul and 2µl was injected in the lateral ventricle by expiratory pressure using a glass micropipette. Following injection, pup brains were placed between electrodes targeted the dorsal SVZ and subjected to five 95V electrical pulses (50 ms, 950 ms intervals). BTX ECM399 (Thermo Fisher Scientific) electroporator was used with 7 mm tweezer electrodes. Electroporated pups were reanimated at 37°C before returning to the mother.

Added in the revised version: In order to analyze the Thbs4 positive cells in the SVZ, we subcloned the eGFP gene (from a pCMV-eGFP N1) into a PGL3-basic backbone containing mouse Thbs4 promoter, kindly provided by Stenina-Adognravi lab (Muppala et al., 2017). pThbs4-GFP was also electroporated in P0-P2 mice, and euthanized 1-day post-electroporation to corroborate Thbs4 expression in the SVZ at postnatal stages.

Added in the revised version: AAV vector production

Added in the revised version: AAV2/9-CAG-HAS2-3xFlag vector was produced by polyethylenimine (PEI) mediated triple transfection of low passage HEK-293 T/17 cells (ATCC; cat number CRL-11268). The respective AAV expression plasmid expression plasmid pNMRHAS2_N1 (provided by Dr. Vera Gorbunova) was cotransfected with the adeno helper pAd Delta F6 plasmid (Penn Vector Core, cat # PL-F-PVADF6) and AAV Rep Cap pAAV2/9 plasmid (Penn Vector Core, cat # PL-TPV008). AAV vector was purified as previously described (Bourdenx et al., 2015). Cells are harvested 72 h post-transfection, resuspended in lysis buffer (150 mM NaCl, 50 mM Tris–HCl pH 8.5) and lysed by 3 freeze-thaw cycles (37°C/-80°C). The cell lysate is treated with 150 units/ml Benzonase (Sigma, St Louis, MO) for 1 h at 37°C, and the crude lysate is clarified by centrifugation. Vectors are purified by iodixanol step gradient centrifugation and concentrated and buffer-exchanged into Lactated Ringer’s solution (Baxter, Deerfield, IL) using vivaspin20 100 kDa cut-off concentrator (Sartorius Stedim, Goettingen, Germany). Titrations were performed at the transcriptome core facility (Neurocentre Magendie, INSERM U862, Bordeaux, France). The genome-containing particle (gcp) titer was determined at a concentration of 1.1013 gcp/ml by quantitative real-time PCR using the Light Cycler 480 SYBR green master mix (Roche, cat # 04887352001) with primers specific for the AAV2 ITRs (fwd 50-GGAACCCCTAGTGATG GAGTT-3’; rev 50-CGGCCTCAGTGAGCGA-30) on a Light Cycler 480 instrument. The purity assessment of vector stocks was estimated by loading 10 μl of vector stock on 10 % SDS acrylamide gels; total proteins were visualized using the Krypton Infrared Protein Stain according to the manufacturer’s instructions (Life Technologies).

Added in the revised version: In vivo AAV-HAS2 injections

Added in the revised version: Mice were anesthetized with isoflurane and fixed to a stereotaxic frame connected with a gas mask. We injected 1µl of adeno-associated virus (AAV2/9) CAG-HAS2-3xFlag (1×10^13 gc/µl) in the contralateral striatum (AP: +1.0; ML: −1.8; DV: −3.2) to simulate scar tissue hyaluronan. After two weeks, the ischemic group was subjected to the MCAO model in the ipsilateral hemisphere and perfused 30 days later. Only slices with representative lesions in both hemispheres were used for image and statistical analysis.

Added in the revised version: Immunofluorescence

Added in the revised version: Mice were anesthetized with an intraperitoneal injection of pentobarbital (100 mg/kg). Brain tissues were fixed by intracardiac perfusion using 4% paraformaldehyde. Fixed brains were sliced at 40 μm of thickness using the Leica VT1200S vibratome (Leica Microsystems) and subjected to immunofluorescence. For standard immunofluorescence protocol, tissue slices were permeabilized with 0.1% Triton and non-specific epitopes blocked with 1% Bovine Serum Albumin and 10% Donkey serum in PBS. Primary antibodies (Supplementary Table 1) were incubated at different concentrations (see Table 1 for specification) overnight at 4°C and then washed three times with 0.1% Triton in PBS. Secondary conjugated antibodies were incubated for 1 h at room temperature. After three washes with 0.1% Triton in PBS, cells were stained for 10 min with DAPI and further washed with PBS. Finally, coverslips were mounted with Fluoromount (SouthernBiotech) and analyzed by fluorescence using the confocal microscope Leica TCS SP8. For hyaluronic acid labeling, endogenous biotins were pre-blocked using Streptavidin/Biotin blocking solution (Vector Labs, #SP-2002) following the manufacturer’s instructions and further blocked with PBS 1X 0.1% Saponin 1% BSA for one hour at room temperature. Primary antibodies and biotinylated Hyaluronic Acid Binding Protein (HABP, #385911 Millipore) were incubated for 72 hours at 4°C or 24 hours at room temperature. After three washes with PBS, 1X 0.1% Saponin, secondary antibody and DAPI diluted in the same blocking buffer were incubated for one hour at room temperature. Slices were mounted with Mowiol (Millipore #475904) in superfrost slides (Thermo Fisher Scientific) and #1.5 coverslips.

Added in the revised version: Imaging and image analysis

Added in the revised version: Confocal images were acquired in a Leica TCS SP8 with 20x (0.8 NA) and 63x objectives (1.4 NA). Images were taken with a pixel size of 90 nm. Widefield imaging was performed in a Zeiss AxioVision fluorescence microscope and a 3DHISTECH panoramic MIDI II digital slide scanner. Settings were adjusted for each individual experiment and kept for all conditions. Image processing was performed using FIJI (Schindelin et al., 2012), CaseViewer and Hyugens (SVI, The Netherlands) software. For Thbs4, Nestin and DCX immunofluorescence, wholemount SVZ preparations and AAV-HAS2 analysis, tile-scan was done using the navigator plug-in in the LasX software. For Nestin immunofluorescence, images were deconvoluted using the Hyugens software.

Added in the revised version: SVZ images were reconstructed using the tile-scan tool (LasX software, Leica). SVZ region of interest (ROI) was delimited manually and expression of each marker was measured in the SVZ ROI. Quantification of all markers was done manually, except for HABP quantification, which was done with custom FIJI scripts for in vitro (https://github.com/MariArdaya/Hyaluronan_v2) and in vivo analysis (https://github.com/SoriaFN/Tools). The HABP signal was segmented with an automatic threshold and colocalized with other markers such as CD44. This, combined with cell segmentation, based on the Thbs4 mask, allowed for quantification of HABP area and number of spots inside cells, outside cells or at the membrane (with a band ROI of 1 µm). The Analyze Skeleton and Fractal count plugins were used on the skeletonized segmented HABP to quantify the longest shortest path and the fractal dimension, respectively, to reveal the length and interconnectivity of the matrix network (Casey et al., 2024; Soria et al., 2020).

Added in the revised version: Western blot

Added in the revised version: The SVZ was freshly dissected on ice and minced using a micro grinder. Total protein was extracted after tissue lysis with RIPA buffer (Thermo Fisher Scientific #89900) and Protease and Phosphatase Inhibitor Cocktail 100X (Thermo Fisher Scientific). Samples were lysed in ice for 10 minutes and sonicated using an Ultrasonic Processor (Hielscher Ultrasound Technology). Sonicated tissues were centrifuged at 1200 rpm for 10 min and 4°C to remove insoluble fragments. Protein concentration was measured by chemioluminescence using RC DC™ Protein Assay (Bio-Rad). Forty micrograms of protein samples were denatured in reducing loading buffer containing β-mercaptoethanol (M3148, Sigma) and 0.0002% bromophenol blue at 95°C for 5 minutes and separated in pre-casted 12% or 7.5% Tris-Glycine polyacrylamide gel using the Criterion cell system (Biorad). Electrophoresis was conducted at 80 V in a Tris-Glycine buffer (25 mM Tris, 192 mM glycine, 0.1% SDS in dH2O, pH 8.3). After complete separation, proteins were transferred to a nitrocellulose membrane (Biorad) using Trans-Blot® TurboTM Transfer System (Bio-Rad). Nitrocellulose membranes were incubated in blocking solution (Tris-Buffer Solution (TBS), 0.1% Tween-20 and 5% BSA) for one hour at room temperature. Primary antibodies were incubated overnight at 4°C. Horseradish peroxidase-conjugated secondary antibodies were incubated in blocking buffer for one hour at room temperature. Immunodetection was performed by electrochemical luminescence in a Chemidoc MP (Biorad) and the intensity of the bands was quantified using ImageLab (Bio-Rad) software corrected by Gaussian curves.

Added in the revised version: qRT-PCR

Added in the revised version: Fresh SVZ tissue from ipsilateral hemisphere was isolated as previously described from sham, 15 and 30 days after MCAO animals to analyze the expression of several genes (Supplementary Table 2). After mechanical trituration, RNA was extracted from the tissue after mechanical trituration following the manufactureŕs protocol (NZYTECH Total RNA Isolation kit protocol (NZYTECH, MB13402). RNA concentration was measured by spectrophotometry using the Nanodrop 2000c spectrophotometer (Thermo-Scientific, USA). Real time polymerase chain reaction (RT-PCR) from 10 ng of RNA was performed using the One-step NZYTECH RT-qPCR Green kit according to manufactureŕs protocol (NZYTECH, MB34302). RT-PCR was normalized using the housekeeping genes hypoxanthine phosphoribosyl-transferase 1 (HPRT1, Qiagen QuantiTect Primer Assay, Barcelona, Spain). Specificity of each amplification was confirmed by melting curve analysis. PCR reactions were performed in a CFX96 Detection System (BioRad, Madrid, Spain). Semi-quantification was performed using the 2-ΔΔCt algorithm. Unsupervised hierarchical distance matrix and distance matrix map was developed using Orange3 software (Python v3.0).

Added in the revised version: Statistical analysis

Added in the revised version: Statistical analysis was done using GraphPad Prism software (version 8.0; GraphPad software). Data are presented as mean ± SEM unless specified otherwise. Statistical analysis was performed using unpaired Student’s t-test for independent conditions and paired Student’s t-test for dependent conditions. In addition, one-way and two-way ANOVA were also performed to compare more than 2 groups among them. Tukey and Dunnett post-hoc tests were used. For HABP skeleton analysis, Kruskal-Wallis analysis was performed to compare cumulative frequency curves in more than 2 groups. P value was considered significant at p < 0.05. Sample number (n), p values and statistical test are indicated in figure legends.